<a href="https://colab.research.google.com/github/ariel-j/python-learning/blob/main/hw2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "ariel jayson"

```

---

# Assignment 2: Applications of modular arithmetics

```python
import math
```

## Question 1: Cryptography

You will implement and analyze two simple cryptographic algorithms — Caesar's and Scytale ciphers.

1. Implement the encoder for Ceasar's cipher: https://en.wikipedia.org/wiki/Caesar_cipher (encode only the letters, leave other characters as-is).



```
# This is formatted as code
```

Q1.1: Implement function `ceasar_cipher`, which will transform regular text into encoded text
<pre>
------------
Input parameters:
- plaintext - string, the text that should be encrypted
- N - int, size of shift
------------
return value:
string, the text after encoding
</pre>

<!-- grade_id: cell-fcc26ea16c7fcc6e, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher(plaintext,N):
    ans = ""
    for char in plaintext:
        if char.isalpha():
            # Determine the new position of the letter in the alphabet after the shift
            oldPosition = ord(char)
            newPosition = (oldPosition - ord('a') + N) % 26 + ord('a')
            ans += chr(newPosition)
        else:
            # Non-alphabetic characters are left as-is
            ans += char

    return ans
ceasar_cipher("abcd", 2)
```

Output:
```
'cdef'
```

<!-- grade_id: cell-910c1caba85f9efe, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipher' function..\n ")
assert type(ceasar_cipher("meep meep",1))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'ceasarCipher' function..
 
good job!
Tests passed. There may be hidden tests too
```

2. Implement a decoder function that can reverse the cypher you've created, and demonstrate it, by encoding and then decoding the implementation (the source code) of the Caesar cipher using Caesar cipher:
  * Encode the Python code
  * Decode the result of encoding
  * Compare the original and the recovered code  - you can add a box where you print the results — they should be the same.

# Q1.2A: Implement function `ceasar_cipher_decode`, which will restore the cyphertext created by `ceasar_cipher` to plaintext

<pre>
------------
Input parameters:
- cyphertext - string, the encoded text that needs to be decyphered
- N - int, size of shift that was used in the encoding
------------
return value:
string, the text after decoding
</pre>

<!-- grade_id: cell-fd561d08a5195124, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher_decode(cypherText,N=1):
  ans = ""
  for char in cypherText:
      if char.isalpha():
          # Determine the original position of the letter in the alphabet before the shift
          oldPosition = ord(char)
          newPosition = (oldPosition - ord('a') - N) % 26 + ord('a')
          ans += chr(newPosition)
      else:
          ans += char

  return ans

```

<!-- grade_id: cell-a3e55a24e9359804, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(ceasar_cipher_decode("ABC",3))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.2B: Implement function `return_encrypted_ceasar_code`, which will return a string that equals the code of `ceasar_cipher` after encryption, with N=2 . Then, implement a function `return_decoded_ceasar_code` which will call `return_encrypted_ceasar_code` and return the decrypted outcome.
<pre>
return_encrypted_ceasar_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre>

<pre>
return_decoded_ceasar_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_ceasar_code()
------------
return value:
string, the source code restored to plaintext
</pre>

<!-- grade_id: cell-0c2fa8c5fcfb1b21, points: 0, grade: false, solution: true -->
```python
import inspect

def return_encrypted_ceasar_code():
  source_code = inspect.getsource(ceasar_cipher)
  encrypted_code = ceasar_cipher(source_code, 2)
  return encrypted_code

def return_decoded_ceasar_code(encrypted_code=return_encrypted_ceasar_code()):
    decoded_code = ceasar_cipher_decode(encrypted_code, 2)
    return decoded_code

```

<!-- grade_id: cell-290384e75f4cb72f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(return_encrypted_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Find a real English word consisting of at least 3 letters which, when encoded using Caesar's cipher, is transformed to another real English word.

Q1.3: Implement function `return_word_and_shift`, which returns a word in English, 3 letters or more in length, and an integer N so that if we apply Ceasar cypher to the word with shift N, we get another word in English.
<pre>
------------
Input parameters:
none
------------
return value:
(word,shift)
- word - string, a word in English
- shift - integer, the size of N needed to encode it so that the cyphertext is also a word in English
</pre>

<!-- grade_id: cell-61c68e5bc86ea006, points: 0, grade: false, solution: true -->
```python

import random
def return_word_and_shift():
    words = [ "grape", "abjurer", "nowhere", "inkier", "purply","fusion", "layout","manful",  "thumbs","primero", "sulphur", "steeds", "tuffet","the","max", "pear","apple", "banana", "kiwi", "mango"]
    word = random.choice(words)
    shift = random.randint(1, 25)
    encoded_word = ceasar_cipher(word, shift)
    if encoded_word.lower() not in words:
        return return_word_and_shift()
    return (word, shift)


```

<!-- grade_id: cell-fcd88242abcbb20b, points: 10, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_int' function..\n ")
found_word=return_word_and_shift()
assert len(found_word)==2, 'function should return tuple with two values'
assert type(found_word[0]) is str, 'the first part of the output should be of type string'
assert type(found_word[1]) is int, 'the secondpart of the output should be of type int'
assert len(found_word[0])>=3, 'the word should be at least 3 letters long'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_int' function..
 
good job!
Tests passed. There may be hidden tests too
```

4. Implement the encoder for the Vigenère cipher: https://en.wikipedia.org/wiki/Vigen%C3%A8re_cipher
(encode only the letters, leave other characters as-is).

**Q1**.4: Implement function `vigenere_encode`, which will encrypt the input text according to the vigenère code with the provided key.
<pre>
------------
Input parameters:
- plaintext - string, the text that needs to be encoded
- key - str
------------
return value:
string, the text after encoding
</pre>

<!-- grade_id: cell-a74f19da20d8c32c, points: 0, grade: false, solution: true -->
```python
def vigenere_encode(plaintext,key):
  plaintext = plaintext.upper()
  key = key.upper()
  ciphertext = ""
  keyIndex = 0
  for char in plaintext:
      if char.isalpha():
          char_value = ord(char) - ord('A')
          key_value = ord(key[keyIndex % len(key)]) - ord('A')
          cipher_value = (char_value + key_value) % 26
          cipher_char = chr(cipher_value + ord('A'))
          ciphertext += cipher_char
          keyIndex += 1
      else:
          ciphertext += char
  return ciphertext

```

<!-- grade_id: cell-fede5aa2fb49e35f, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_encode' function..\n ")
assert type(vigenere_encode("ABCabcABC","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'vigenere_encode' function..
 
good job!
Tests passed. There may be hidden tests too
```


5. Implement a decoder, and demonstrate by encoding and then decoding the implementation of the Vigenère cipher, as before.

<b>Please note:</b> in the case where the encoded text length isn't a multiple of the key, the decoder should still be able to reverse the encoding correctly! Make sure you've handled this case.

Q1.5A: Implement function `vigenere_decode`, which will restore the cyphertext created by `vigenere_encode` to plaintext
<pre>
------------
Input parameters:
- ciphertext - string, the encoded text that needs to be decrypted
- key - str
------------
return value:
string, the text after decoding
</pre>

<!-- grade_id: cell-1780c42a6fe02cf5, points: 0, grade: false, solution: true -->
```python
def vigenere_decode(ciphertext,key):
  ans = ""
  ciphertext = ciphertext.upper()
  key = key.upper()
  ciphertext = ''
  keylen = len(key)
  keyIndex = 0
  for i in range(len(ciphertext)):
      if ciphertext[i].isalpha():
          shift = ord(key[keyIndex]) - ord('A')
          ciphertext += chr((ord(ciphertext[i]) - ord('A') + shift) % 26 + ord('A'))
          keyIndex = (keyIndex + 1) % keylen
      else:
          ciphertext += ciphertext[i]
  return ciphertext
```

<!-- grade_id: cell-7525afb85c29fb34, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_decode' function..\n ")
assert type(vigenere_decode("ABCDefgh","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'vigenere_decode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.5B: Implement functions `return_encrypted_vigenere_code`, which will return a string that equals the code of `vigenere_encode` after encryption, with key="vig". Then, implement a function `return_decoded_vigenere_code` which will call `return_encrypted_vigenere_code` and return the decrypted outcome. You may print the output to verify that it matches the original.
<pre>
return_encrypted_vigenere_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre>

<pre>
return_decoded_vigenere_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_vigenere_code()
------------
return value:
string, the source code restored to plaintext
</pre>

<!-- grade_id: cell-0e04f341a44584a7, points: 0, grade: false, solution: true -->
```python

def return_encrypted_vigenere_code():
    source_code = inspect.getsource(vigenere_encode)
    encrypted_code = vigenere_encode(source_code, 'vig')
    return encrypted_code


def return_decoded_vigenere_code(encrypted_code=return_encrypted_vigenere_code()):
    decoded_code = vigenere_decode(encrypted_code, 'vig')
    return decoded_code

```

<!-- grade_id: cell-bef06a38f2c90ce7, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'return_encrypted/return_decoded' function..\n ")
assert type(return_encrypted_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'return_encrypted/return_decoded' function..
 
good job!
Tests passed. There may be hidden tests too
```

# Question 2: Hash tables

The code cell below contains emergency phone numbers of the university (taken from https://in.bgu.ac.il/logistics/Pages/dialing_rules.aspx). A single entity may have several phone numbers.

You will implement a ‘truecaller’ facility: given a phone number, output the name of the entity. For example,
```python
caller = truecaller(t, 4100)
print(f"{caller} called!")
```
```truecaller``` should return “Police”.

if the number is not found, the function should return None

Q1. Implement 'truecaller' using a Python dictionary


```python
phonebook = [["Security",     '61888', '61555', '61444'],
    ["Police",    '100', '4100'],
    ["Magen David Adom",    '101', '4101'],
    ["Fire brigade",    '102', '4102'],
    ["Soroka hospital",    '4007'] ]

def truecaller(phonebook, number):
    phone_dict = {}
    for entry in phonebook:
        entity = entry[0]
        for phone in entry[1:]:
            phone_dict[phone] = entity
    if number in phone_dict:
        return phone_dict[number]
    else:
        return None


```

Q2.1A : Implement an auxiliary `make_trucaller_python_dictionary` function that converts the provided `phonebook` list into a Python dictionary
<pre>
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
Python dictionary
------------
</pre>

<!-- grade_id: cell-87a450fdc22b0d47, points: 0, grade: false, solution: true -->
```python
def make_truecaller_python_dictionary(phn=phonebook):
    Dictionary = {}
    for entry in phn:
        Name = entry[0]
        Numbers = entry[1:]
        Dictionary[Name] = Numbers
    return Dictionary

```

Q2.1B : Implement a `truecaller1` function, relying on `make_truecaller_python_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre>
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-9e86c1484dcf36f4, points: 0, grade: false, solution: true -->
```python
phonebookDic= make_truecaller_python_dictionary(phonebook)
def truecaller1(phone_number,caller_dict=phonebookDic):
   for Name, phones in phonebookDic.items():
        if phone_number in phones:
          return Name
   return None


```

<!-- grade_id: cell-1be6975be3ec4742, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_python_dictionary())==dict,'output type is incorrect, should be python dictionary'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller1('101'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


Q2. Implement "truecaller2" using your own implementation of dictionary:





<!-- grade_id: cell-f06f7223c2e94a50, points: 0, grade: false, solution: true -->
```python
#we can use a similar dictionary as we saw in class
def maket(n):
    # t = []
    # for _ in range(n):
    #     t.append([])
    # return t
    return [list() for _ in range(n)]

def tset(table, key, value, h=hash):
    i = h(key) % len(table)
    entry = table[i]
    for i, (k, v) in enumerate(entry):
        if key==k:
            entry[i] = (key, value)
            return
    entry.append((key, value))

def tget(table, key, h=hash):
    i = h(key) % len(table)
    entry = table[i]
    for i, (k, v) in enumerate(entry):
        if key==k:
            return v
    return None

```

Q2.2A : Implement an auxiliary `make_truecaller_own_implement_dictionary` function that converts the provided `phonebook` list into a structure that behaves like a hash table
    
<pre>
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
a Python list
------------
</pre>

<!-- grade_id: cell-0f5cd96a463d46c9, points: 0, grade: false, solution: true -->
```python
def make_truecaller_own_implement_dictionary(phBook=phonebook):
   phonebookDic = [[] for _ in range(1000)]
   for entry in phBook:
        callerName = entry[0]
        phoneNumbers = entry[1:]
        for phoneNumber in phoneNumbers:
            hashValue = hash(phoneNumber) % 1000
            phonebookDic[hashValue].append((phoneNumber,callerName))
   return phonebookDic

```

Q2.2B : Implement a `truecaller2` function, relying on `make_truecaller_own_implement_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre>
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-a2c907715447ebae, points: 0, grade: false, solution: true -->
```python
make_own_dictionary=make_truecaller_own_implement_dictionary(phonebook)
def truecaller2(phone_number,own_dictionary=make_own_dictionary):
    hashValue = hash(phone_number) % 1000
    for key, value in make_own_dictionary[hashValue]:
        if key == phone_number:
            return value
    return None

```

<!-- grade_id: cell-9bd7a35c7e4305a6, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_own_implement_dictionary())==list,'output type is incorrect, should be list'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller2('102'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Implement a hash function `hash_string` for ids. Treat ids as _strings_ of characters.


<pre>
------------
Input parameters:
- str_to_hash - a string
------------
return value:
an integer
------------
</pre>

<!-- grade_id: cell-907e0df43371898d, points: 0, grade: false, solution: true -->
```python
def hash_string(str_to_hash):
  hash_value = 0
  for char in str_to_hash:
    hash_value = (hash_value * 31 + ord(char)) % 2**32
  return hash_value
```

<!-- grade_id: cell-6074fd7cee4c2b6d, points: 12, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_string' function..\n ")
assert type(hash_string("some string"))==int,'output should be an integer'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_string' function..
 
good job!
Tests passed. There may be hidden tests too
```


4. A check digit is used in ids in order to detect errors and make faking an id a bit harder. The exact formula for ids' check digit can be found in https://he.wikipedia.org/wiki/%D7%A1%D7%A4%D7%A8%D7%AA_%D7%91%D7%99%D7%A7%D7%95%D7%A8%D7%AA. Implement a function `check_id` which checks if a given id is legal with 2 conditions:
1. It is of length 9
2. The check digit is correct

Implement the function `check_id` that returns whether an id is legal or not:





<pre>
------------
Input parameters:
- id: a string
------------
return value:
- boolean
------------
</pre>

<!-- grade_id: cell-a5389da6f18d38bd, points: 0, grade: false, solution: true -->
```python
def check_id(id):
    # Check length
    if len(id) != 9:
        return False

    # Calculate check digit
    factors = [1, 2, 1, 2, 1, 2, 1, 2, 1]
    digits = [int(d) for d in id]
    products = [factors[i] * digits[i] for i in range(8)]
    products = [p % 10 + p // 10 for p in products]
    check_digit = (10 - sum(products)) % 10

    # Check check digit
    return check_digit == int(id[8])

```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
```python


# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the min_size- functions..\n ")
assert check_id("123456789")==False,'output should be False'
assert check_id("123456782")==True,'output should be True'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the min_size- functions..
 
good job!
Tests passed. There may be hidden tests too
```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
5. Hashing is called ‘perfect’ if there are no ‘key conflicts‘. Can you find if there is a value that has a conflict with your own id, when using the hash function you implemented in Q2.3?

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
Implement the function `find_conflict` that returns if there is an id with the same hash value as your own id, or an empty string otherwise:

<pre>
------------
Input parameters:
- id: your own id
------------
return value:
- string, either empty or a legal id with the same hash value as your own id
------------
</pre>

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python

def find_conflict(id):
    hash_value = hash_string(id)
    for i in range(1000000000):
        test_id = str(i).zfill(9)
        if hash_string(test_id) == hash_value and check_id(test_id):
            return test_id
    return ''

```