<a href="https://colab.research.google.com/github/ariel-j/python-learning/blob/main/hw3_.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 3: Complex numbers and Fourier transform
### Due Date (22.5 - 23:59)

<a id="gen_instructions"></a>
## Working on the assignment, todo and not todo:
<b><u>Todo</u></b>:<br />
<div class="alert alert-success">
<ul><li>&#x2714; <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>&#x2714; Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>&#x2714; to <b>submit</b> the assignment, download the notebook (File → Download .ipynb in Google Colab)</li>
<li>&#x2714; submit <b>only</b> the <b>ipynb</b> file under the name <b>hw3.ipynb</b></li>
<li>&#x2714; It is advisable to add extra cells to check your code implementation</li></ul>
</div>
<b><u>Not todo</u></b>:<br />
<div class="alert alert-danger">
<ul><li>&#x2718; do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>&#x2718; do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>&#x2718; <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>&#x2718; <b>do NOT change</b> the notebook file name</li></ul>
</div>
<b><u>Not good - Any of the following</u></b>:<br/>
<div class="alert alert-danger">
<ul><li>&#x2718; Code not Implemented</li>
<li>&#x2718; No returned value</li>
<li>&#x2718; Did not remove the '<code>raise NotImplementedError()</code>' line</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
</code></div>


```python
# import librarys :
import matplotlib.pyplot as plt
import numpy as np
import numpy.testing as npt
from scipy import fftpack
from matplotlib.colors import LogNorm
from unittest.mock import patch
np.random.seed(42)
```

## Question 1: Complex exponent

Complex exponent is the complex version of the library function `exp`.

<div class="alert alert-danger">
Throughout this question, when you are asked to implement a formula, you must implement the exact formula as written in the question, and you may not use the built in operations of python libraries!
</div>


Q1.1 Implement complex exponent using real-valued functions from the `numpy` library. Use Euler's formula for the exponent complex number, as in the formula below<br><br>
**Euler's formula** $$e^{a+ib}=e^a⋅(cosb+isinb)$$
<pre>
compute complex_exponent(z)
------------
Input parameters:
- z - complex number (a+bi)
------------
return value:
- cexp - complex exponent
------------
</pre>

<!-- grade_id: cell-1f2faa78562b319c, points: 0, grade: false, solution: true -->
```python
import math
def complex_exponent(z):
    a = z.real
    b= z.imag
    Exp_Real = np.exp(a) * np.cos(b)
    Exp_Imag = np.exp(a) * np.sin(b)
    Cexp = Exp_Real + Exp_Imag * 1j
    return Cexp


```

<!-- grade_id: cell-d7caf9133b0580d2, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print ("Test - Testing the implementation of the 'complex_exponent' function..\n ")
z = 0+0j
print(f'complex number z = {z}')
print(f"your real number is z.real = {z.real}")
print(f"your imagnary number is z.imag = {z.imag}")
print(f"complex euler of z is ℯᶻ = {complex_exponent(z)}")
assert 1+0j == complex_exponent(z)


print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

Output:
```
Test - Testing the implementation of the 'complex_exponent' function..
 
complex number z = 0j
your real number is z.real = 0.0
your imagnary number is z.imag = 0.0
complex euler of z is ℯᶻ = (1+0j)
good job!
Sanity tests passed. There are additional hidden tests...
```

Q1.2 Implement addition operation on two complex numbers as in the formula below .<br><br>
**complex addition formula** $$(a + bi) + (c + di) = (a + c) + (b + d)i$$
<pre>
compute complex_addition(z1,z2)
------------
Input parameters:
- z1 - complex number (a + bi)
- z2 - complex number (c + di)
------------
return value:
- z - addition result complex number
------------
</pre>


<!-- grade_id: cell-95fea2ef23e5280a, points: 0, grade: false, solution: true -->
```python
def complex_addition(z1,z2):
    Real = z1.real + z2.real
    Imag = z1.imag + z2.imag
    z = Real + Imag * 1j
    return z

```

<!-- grade_id: cell-44a26f249a2ef03f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2 --- Test your implementation:
# ---------------------------
print ("Test - Testing the implementation of the 'complex_addition' function..\n ")
z1 = 0+1j
z2 = 1+0j
print(f'complex number z1 = {z1}')
print(f'complex number z2 = {z2}')
print(f'your addition result number is  = {complex_addition(z1,z2)}')

assert 1+1j == complex_addition(z1,z2)

print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

Output:
```
Test - Testing the implementation of the 'complex_addition' function..
 
complex number z1 = 1j
complex number z2 = (1+0j)
your addition result number is  = (1+1j)
good job!
Sanity tests passed. There are additional hidden tests...
```

Q1.3 Implement convert_to_polar on a complex number which given a complex number of the cartesian form $$z=a+bi$$ converts it to its polar form $$z=r\cdot e^{i\theta}$$
You can use np.abosulte and np.angle
<pre>
compute convert_to_polar(z)
------------
Input parameters:
- z - complex number (a + bi)
------------
return value:
- z_polar - polar form of z
------------
</pre>

```python
def convert_to_polar(z):
    r = np.abs(z)
    theta = np.angle(z)
    Polar_z = r * np.exp(1j * theta)
    return Polar_z
```

```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.3 --- Test your implementation:
# ---------------------------
print ("Test - Testing the implementation of the 'convert_to_polar' function..\n ")
z = 1+1j
print(f'complex number z = {z}')
print(f'your polar form is  = {convert_to_polar(z)}')

print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

Output:
```
Test - Testing the implementation of the 'convert_to_polar' function..
 
complex number z = (1+1j)
your polar form is  = (1.0000000000000002+1j)
good job!
Sanity tests passed. There are additional hidden tests...
```

Q1.4 Implement division operation on two complex numbers as in the formula below. guidance: you can use np.absolute() and np.angle() <br><br>
**complex division formula** $$ \frac{r_{1}\cdot e^{i\theta_{1}}}{r_{2}\cdot e^{i\theta_{2}}} = \frac{r_{1}}{r_{2}} ⋅ e^{i(\theta_{1}-\theta_{2})} $$
<pre>
compute complex_div(z1,z2)
------------
Input parameters:
- z1 - complex number (r1 * exp(θ1⋅i))
- z2 - complex number (r2 * exp(θ2⋅i))
------------
return value:
- z - division result complex number
------------
</pre>


<!-- grade_id: cell-57c68f6c686fa47a, points: 0, grade: false, solution: true -->
```python
def complex_div(z1,z2):
    r1 = np.abs(z1)
    theta1 = np.angle(z1)
    r2 = np.abs(z2)
    theta2 = np.angle(z2)
    r = r1 / r2
    theta = theta1 - theta2
    z = r * np.exp(1j * theta)
    return z

```

<!-- grade_id: cell-710c3582c36821f4, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4 --- Test your implementation:
# ---------------------------
print ("Test - Testing the implementation of the 'complex_addition' function..\n ")
z1 = 1+1j
z2 = 1+1j
print(f'complex number z1 = {z1}')
print(f'complex number z2 = {z2}')
print(f'your division result number is  = {complex_div(z1,z2)}')
print(f'if your number is close enough to the real result it\'s fine')

npt.assert_almost_equal( z1/z2 , complex_div(z1,z2) )


print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

Q1.5 Implement Conjugate operation on complex number without using built in np.conj function
<pre>
compute conjug(z)
------------
Input parameters:
- z - complex number (a + bi)
------------
return value:
- z - Conjugate operation result as complex number
------------
</pre>


<!-- grade_id: cell-4419cea85136f070, points: 0, grade: false, solution: true -->
```python
def conjug(z):
    conjug_z = z.real - z.imag * 1j
    return conjug_z
```

<!-- grade_id: cell-4d1520199c15506e, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5 --- Test your implementation:
# ---------------------------
print ("Test - Testing the implementation of the 'conjug' function..\n ")
z = 1+1j
print(f'complex number z = {z}')
print(f'complex number z-conjugate = {conjug(z)}')
assert isinstance(conjug(z),complex) ,'wrong type'
print ("good job!\nSanity tests passed. There are additional hidden tests...")

```

## Question 2: Image processing with Fourier transform

```python
import PIL
import urllib

### Below you'll find a sample greyscale image
pil_img = PIL.Image.open(urllib.request.urlopen('http://palmzlib.sourceforge.net/images/boat1.png'))
img = np.asarray(pil_img)
orig_img = img.copy()

plt.figure()
plt.imshow(orig_img, plt.cm.gray)
plt.title('Original image')
```

Q2.1. Add $\mathcal{N}$(0,20) noise to the given image

Implement the function add_noise:

<pre>
Input parameters:
im - an image
noise - the noise to add

return value - the image with noise
</pre>

Display the image with the noise given in the code below (which generates a normal distribution of noise)
<pre>
noise = np.random.normal(0.0, 20, size=img.shape)
</pre>

<!-- grade_id: cell-bd1d883ea75be84d, points: 0, grade: false, solution: true -->
```python
noise = np.random.normal(0.0, 20.0, size=img.shape)

def add_noise(im, noise):
    Noisy_Im = im.astype(float) + noise
    return Noisy_Im

```

<!-- grade_id: cell-781935dce6327fd4, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'low_pass_filter' function..\n ")
noise = np.random.normal(0.0, 20, size=img.shape)
img_noise_added = add_noise(orig_img, noise)
assert orig_img.shape == img_noise_added.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")



```

```python
def plot_spectrum(im_fft):
    plt.imshow(np.abs(im_fft), norm=LogNorm(vmin=5))
    plt.colorbar()
```

Q2.2. Apply 20% low-pass filter to the rows of the image (keep 20% of the frequencies).

Implement the function low_pass_filter:

<pre>
Input parameters:
im - an image
keep_fraction - what fraction of frequencies to keep

return value - the image after low-pass filter was applied
</pre>

Display the filtered image.

<!-- grade_id: cell-48cb1f0a772f8f43, points: 0, grade: false, solution: true -->
```python

def low_pass_filter(im, keep_fraction):
    DFT = np.fft.fft2(im)
    DFT_Shifted = np.fft.fftshift(DFT)
    Magnitude_Spectrum = np.abs(DFT_Shifted)
    rows, cols = im.shape
    cy = rows/2
    cx =  cols/2
    Cutoff_Radius = np.sqrt(cx**2 + cy**2) * np.sqrt(1 - keep_fraction)
    for i in range(rows):
        for j in range(cols):
            if np.sqrt((i-cy)**2 + (j-cx)**2) > Cutoff_Radius:
                DFT_Shifted[i,j] = 0
    DFT_Backshifted = np.fft.ifftshift(DFT_Shifted)
    Im_Back = np.fft.ifft2(DFT_Backshifted)
    return np.real(Im_Back)

```

<!-- grade_id: cell-71ac68e593a09556, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'low_pass_filter' function..\n ")
row_fraction = 0.2
img_new = low_pass_filter(img_noise_added, row_fraction)
assert orig_img.shape == img_new.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")



```

Q2.3. Apply 10% high-pass filter to the columns of the image (keep 10% of the frequencies).

Implement the function high_pass_filter:

<pre>
Input parameters:
im - an image
keep_fraction - what fraction of frequencies to keep

return value - the image after high-pass filter was applied
</pre>

Display the filtered image.

<!-- grade_id: cell-b3574377d906bc90, points: 0, grade: false, solution: true -->
```python
def high_pass_filter(im, keep_fraction):
    Im_Fft = fftpack.fft2(im)
    Im_Fft2 = Im_Fft.copy()
    f = Im_Fft2.shape[1]
    Im_Fft2[:, int(f*keep_fraction):int(f*(1-keep_fraction))] = 0
    return fftpack.ifft2(Im_Fft2).real
```

<!-- grade_id: cell-73047ce23f36a781, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'high_pass_filter' function..\n ")
col_fraction = 0.1
img_new = high_pass_filter(img_noise_added, col_fraction)
assert orig_img.shape == img_new.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")
plt.imshow(img_new, plt.cm.gray)
plt.show()




```

Q2.4. Denoise the image using FFT (similarly to the method shown in class).

Implement the function denoise

<pre>
Input parameters:
im - an image
low_pass_fraction - what fraction of low frequencies to keep
high_pass_fraction - what fraction of high frequencies to keep

return value - the denoised image
</pre>

Display the denoised image - apply 20% low-pass filter and 10% high-pass filter.

<!-- grade_id: cell-0797f4873daf1efd, points: 0, grade: false, solution: true -->
```python
def denoise(im, low_pass_fraction, high_pass_fraction):

  return high_pass_filter(low_pass_filter(im, low_pass_fraction), high_pass_fraction)

def Denoise(im, low_pass_fraction, high_pass_fraction):
    Low_Passed_im = low_pass_filter(im, low_pass_fraction)
    high_pass_filter = high_pass_filter(low_pass_filter(im, low_pass_fraction), high_pass_fraction)
    return high_pass_filter

plot_spectrum(denoise(img_noise_added, 0.2, 0.1))
plt.figure()
plt.imshow(denoise(img_noise_added, 0.2, 0.1), plt.cm.gray)
```

<!-- grade_id: cell-c86f92ce7c9f061b, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'denoise' function..\n ")
img_new = denoise(img_noise_added, 0.2, 0.1)
assert orig_img.shape == img_new.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")
plt.imshow(img_new, plt.cm.gray)
plt.show()


```

Q2.5. Calculate the noise removed using FFT

Implement the function noise_removed:

<pre>
Input parameters:
im - an image
low_pass_fraction - what fraction of low frequencies to keep
high_pass_fraction - what fraction of high frequencies to keep

return value - the noise removed
</pre>

Display the noise removed

<!-- grade_id: cell-faf23d8fb3562144, points: 0, grade: false, solution: true -->
```python

def noise_removed(im, low_pass_fraction, high_pass_fraction):
    Denoised_Im = denoise(im, low_pass_fraction, high_pass_fraction)
    Noise_Removed_Im = im - Denoised_Im
    return Noise_Removed_Im
```

<!-- grade_id: cell-2df9ce14b2905d46, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.5 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'denoise' function..\n ")
img_noise_removed = noise_removed(img_noise_added, 0.2, 0.1)
assert orig_img.shape == img_noise_removed.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")
plt.imshow(img_new, plt.cm.gray)
plt.show()


```

Q2.6. Reconstruct the image with original image with noise using the two images calculated in 5.4 and 5.5

Implement the function recover_image:

<pre>
Input parameters:
denoised_im - the image after noise was removed
noise_im - the noise removed
return value - the original image with noise recovered by this function
</pre>

Display the recovered image

<!-- grade_id: cell-96933b85f8658d3a, points: 0, grade: false, solution: true -->
```python
def recover_image(denoised_im, noise_im):
    Recovered_Im = denoised_im + noise_im
    return Recovered_Im

denos = denoise(orig_img, 0.2, 0.1)
nois = noise_removed(img_noise_added, 0.2, 0.1)
plot_spectrum(recover_image(denos, nois))
plt.figure()
plt.imshow(recover_image(denos, nois), plt.cm.gray)
```

<!-- grade_id: cell-bdd8c8e3fd8e40b9, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.6 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'denoise' function..\n ")
denoised_im = denoise(img_noise_added, 0.2, 0.1)
noise_im = noise_removed(img_noise_added, 0.2, 0.1)
img_recovered = recover_image(denoised_im, noise_im)
assert orig_img.shape == img_recovered.shape
print ("good job!\nSanity tests passed. There are additional hidden tests...")
plt.imshow(img_new, plt.cm.gray)
plt.show()


```

Q2.7. Repeat the last question with different keep_fractions values (the parameter to the apply filters functions). Which gives the best results? Use at least 5 different values

Implement the function recover_best:

<pre>
Input parameters:
im - the image with noise
keep_fractions - a list of tuples (low,high) of different keep_fractions values
return value - the keep_fraction value which returns the best results
</pre>

Display the following for each tuple:
1. The error
2. Display the recovered image

At the end, display the best result

```python
def recover_best(im, keep_fractions):
    low_Err = np.inf
    Best_Fil_Frac,  Best_Filtered = None

    for fraction in keep_fractions:
      Filtered = denoise(im, fraction[0], fraction[1])
      med = np.mean((im - Filtered) ** 2)
      if med < low_Err:
        Best_Filtered = Filtered
        Best_Fil_Frac = fraction
        lowest_error = med

    plt.figure()
    plt.imshow(Best_Filtered, plt.cm.gray)
    plt.title(Best_Fil_Frac)
    return Best_Fil_Frac
```

```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.6 --- Test your implementation:
# ---------------------------
print("Testing the implementation of the 'denoise' function..\n ")
best_fraction = recover_best(img_noise_added, [(0.1,0.1), (0.1,0.2), (0.4,0.3)])
print ("good job!\nSanity tests passed. There are additional hidden tests...")
```